---
  title: "Теоретическая вставка"
  description: "This is a template chapter."
  v2: true

---
## Ex 1.1

```yaml
type: NormalExercise
lang: r
xp: 100
skills: 1
key: 91681de4cc



```

This is the assignment text. It should help provide students with the background information needed.
The instructions that follow should be in bullet point form with clear guidance for what is expected.

`@instructions`
- Instruction 1
- Instruction 2
- Instruction 3
- Instruction 4

`@hint`
- Here is the hint for this setup problem. It should get students 50% of the way to the correct answer.

`@pre_exercise_code`
```{r}
# Load datasets and packages here.
```
`@sample_code`
```{r}
# Your
# sample
# code
# should
# be
# ideally
# 10 lines or less,
# with a max
# of 16 lines.
```
`@solution`
```{r}
# Answer goes here
# Make sure to match the comments with your sample code
# to help students see the differences from solution
# to given.
```
`@sct`
```{r}
# Update this to something more informative.
success_msg("Some praise! Then reinforce a learning objective from the exercise.")
```





---
## Лекция о сибах

```yaml
type: VideoExercise
lang: r
xp: 50

key: 45bf51b8f6
video_link: //player.vimeo.com/video/160339466


```

`@projector_key`
31ed9d9b7b89c18699d6c79ab7a10293

---
## Insert exercise title here

```yaml
type: VideoExercise

xp: 50

key: b7602dbf6e



```

`@projector_key`
1971e1d018fb2c898d55d68eee8657d1